Before you turn this problem in, make sure everything runs as expected. First, **restart the kernel** (in the menubar, select Kernel$\rightarrow$Restart) and then **run all cells** (in the menubar, select Cell$\rightarrow$Run All) to avoid typical problems with Jupyter notebooks. **Unfortunately, this does not work with Chrome right now, you will also need to reload the tab in Chrome afterwards**.

Make sure you fill in any place that says `YOUR CODE HERE` or "YOUR ANSWER HERE". Please put your name here:

```python
NAME = "AVISHA ANILKUMAR BHIRYANI"
```

---

# Latent Semantic Indexing

First we will implement latent semantic indexing (LSI)

```python
### Load the input data - do not modify
import json, gzip, numpy as np
raw = json.load(gzip.open("/data/simpsonswiki.json.gz", "rt", encoding="utf-8"))
titles, texts, classes = [x["title"] for x in raw], [x["text"] for x in raw], [x["c"] for x in raw]
```

```python
### Vectorize the text - do not modify
from sklearn.feature_extraction.text import TfidfVectorizer, CountVectorizer
vect = TfidfVectorizer(stop_words="english", sublinear_tf=True, smooth_idf=False, min_df=5)
vect.fit(texts)
vect.idf_ -= 1
idf = vect.idf_
tfidf = vect.transform(texts)
vocabulary = vect.get_feature_names_out()
```

## Implement LSI

Implement Latent Semantic Indexing. Do **not** use regular SVD, but instead use truncated SVD from sklearn. (Do not attempt to implement Truncated SVD yourself, use the library here.) Return weights how well the factors explain the topics.

<!-- grade_id: cell-8eb426fb1c1146b0, points: 0, grade: false, solution: true -->
```python
# Implement LSI here
def lsi(tfidf, k):
    """Latent Semantic Indexing. Return the factors, document assignment, and factor weights"""
    from sklearn.decomposition import TruncatedSVD
    # YOUR CODE HERE
    svd = TruncatedSVD(n_components=k)
    svd.fit(tfidf)
    #V matrix - column wise
    factors = svd.components_
    #U matrix - row vise
    assignment = svd.transform(tfidf)
    weights = svd.explained_variance_
    return factors, assignment, weights
```

```python
_tmp = lsi(tfidf, 2)
print(_tmp[0].shape)
print(_tmp[1].shape)
```

Output:
```
(2, 14153)
(10126, 2)
```

<!-- grade_id: cell-30066761f95116e7, points: 1, grade: true, solution: false -->
```python
### Automatic tests. You do not need to understand or modify this code.
_tmp = lsi(tfidf, 2)
assert len(_tmp) == 3, "Incomplete result"
assert _tmp[0].shape == (2, tfidf.shape[1]), "Factor shape is not correct."
assert _tmp[1].shape == (tfidf.shape[0], 2), "Assignment shape is not correct."
del _tmp
```

## Explore your result

Explore the result: write a function to determine the most important words for each factor, and the most relevant documents.

<!-- grade_id: cell-c3a23922942e74ed, points: 0, grade: false, solution: true -->
```python
def most_important(vocabulary, factor, k=10):
    """Most important words for each factor"""
    # YOUR CODE HERE
    dict_words = {}
    i = 0
    for key in vocabulary:
        dict_words[key] = factor[i]
        i = i+1
    sorted_words = dict(sorted(dict_words.items(), key=lambda item: item[1], reverse=True))
    arr = [each for each in list(sorted_words)[:k]]
    return arr
        
def most_relevant(assignment, k=5):
    """Most relevant documents for each factor (return document indexes)"""
    # YOUR CODE HERE
    arr = []
    dict_docs = {}
    for i in range(0,len(assignment)):
        dict_docs[i] = assignment[i]
        i = i+1
    sorted_docs = dict(sorted(dict_docs.items(), key=lambda item: item[1], reverse=True))
    arr= [each for each in list(sorted_docs)[:k]]
    return arr

def explain(vocabulary, titles, classes, factors, assignment, weights=None):
    """Print an explanation for each factor.
       If weights is None, use the relative share of the assignment weights.
       Print the ARI when assigning each document to its maximum only."""
    from sklearn.metrics import adjusted_rand_score
    # YOUR CODE HERE
    if weights is not None:
        total_factors = factors.shape[0]
        for i in range(0,total_factors):
            important_words = most_important(vocabulary, factors[i])
            important_docs = most_relevant(assignment[:i])
            print(i)
            print(important_docs)
            print(important_words)
        
```

```python
explain(vocabulary, titles, classes, *_tmp)
```

Output:
```
0
[]
['couch', 'simpson', 'homer', 'bart', 'lisa', 'marge', 'gag', 'family', 'song', 'maggie']
1
[0]
['couch', 'gag', 'simpson', 'plot', 'maggie', 'appearances', 'character', 'season', 'family', 'sit']
```

<!-- grade_id: cell-53f1929d10c5494e, points: 1, grade: true, solution: false -->
```python
### Automatic tests. You do not need to understand or modify this code.
_tmp = lsi(tfidf, 2)
assert len(most_important(vocabulary, _tmp[0][0], 42)) == 42, "Wrong number of most important words"
for x in most_important(vocabulary, _tmp[0][0]): assert isinstance(x, str), "Most important words are not words"
assert len(most_relevant(_tmp[1][:,0], 42)) == 42, "Wrong number of relevant results."
from unittest.mock import patch
with patch('__main__.most_important') as mock_1, patch('__main__.most_relevant') as mock_2, patch('__main__.print') as mock_3:
    explain(vocabulary, titles, classes, *_tmp)
    assert mock_1.called, "You did not use most_important"
    assert mock_2.called, "You did not use most_central"
    assert mock_3.called, "You did not print"
```

```python
# Explore your result. These should mostly be meaningful topics!
lsi_factors, lsi_assignment, lsi_weights = lsi(tfidf, 6)
explain(vocabulary, titles, classes, lsi_factors, lsi_assignment, lsi_weights)
```

Output:
```
0
[]
['couch', 'simpson', 'homer', 'bart', 'lisa', 'marge', 'gag', 'family', 'song', 'maggie']
1
[0]
['couch', 'gag', 'simpson', 'plot', 'maggie', 'appearances', 'character', 'season', 'family', 'sit']
Input In [122], in explain(vocabulary, titles, classes, factors, assignment, weights)
     33 for i in range(0,total_factors):
     34     important_words = most_important(vocabulary, factors[i])
---> 35     important_docs = most_relevant(assignment[:i])
     36     print(i)
     37     print(important_docs)

Input In [122], in most_relevant(assignment, k)
     19     dict_docs[i] = assignment[i]
     20     i = i+1
---> 21 sorted_docs = dict(sorted(dict_docs.items(), key=lambda item: item[1], reverse=True))
     22 arr= [each for each in list(sorted_docs)[:k]]
     23 return arr

ValueError: The truth value of an array with more than one element is ambiguous. Use a.any() or a.all()
```

<!-- grade_id: cell-cf8643ed125237fa, points: 1, grade: true, solution: false -->
```python
### This cell contains additional tests. You do not need to modify this cell.
```

```python

```